# Module 5 Assignment

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

```python
import pandas as pd
import numpy as np
import random
import nltk
import helper

from sklearn.preprocessing import LabelEncoder
from sklearn.model_selection import train_test_split
from sklearn.feature_extraction.text import TfidfVectorizer, CountVectorizer
from sklearn.naive_bayes import MultinomialNB
from sklearn import metrics
from sklearn.pipeline import Pipeline
from sklearn.metrics import accuracy_score

from nose.tools import assert_equal, assert_is_instance
from nltk.corpus import reuters
from operator import itemgetter

```

For the following problems we will use the NLTK Reuters corpus which contains 10,788 documents which have been classified into 90 topics.

```python
#Let's take a look at the categories
print(reuters.categories())
```

# Problem 1: Find the number of times a particular token occurs

Write a function called $\texttt{token_counter}$ which takes in an nltk corpora and outputs the number of times a particular token occurs in the corpora.

<!-- grade_id: cell-d4609143baf90e44, points: 0, grade: false, solution: true -->
```python
def token_counter(token,corpora):
    '''
    Inputs
    -------
    
    token: the token to search for
    corpora: a list of documents
    
    Returns
    --------
    
    count: the number of times the token occurs in the corpora
    '''
    # YOUR CODE HERE
    
    return count
```

<!-- grade_id: cell-6fc4830ca34cb389, points: 1, grade: true, solution: false -->
```python
num_to = token_counter('to',reuters)
assert_equal(num_to,34035)
num_ship = token_counter('ship',reuters)
assert_equal(num_ship,116)
```

```python
cv = CountVectorizer(analyzer='word', lowercase=True)
```

# Problem 2: Use CountVectorizer to retrieve the corpus vocabulary

Write a function called $\texttt{most_common_words}$ which uses CountVectorizer to return the vocabulary of the corpora.

<!-- grade_id: cell-320c4d5d381fae3f, points: 0, grade: false, solution: true -->
```python
def most_common_words(corpora):
    '''
    Inputs
    -------

    corpora: a list of documents
    
    Returns
    --------
    
    vocabulary: a Python dictionary containing each word and its count
    '''
    # YOUR CODE HERE
    
    return vocabulary
```

<!-- grade_id: cell-d6a90ab9d83a28cf, points: 1, grade: true, solution: false -->
```python
vocabulary = most_common_words(reuters)
assert_is_instance(vocabulary,dict)
assert_equal(vocabulary['asian'],3384)
```

# Problem 3: Reading and Preprocessing Data
Read in the [badges dataset](https://archive.ics.uci.edu/ml/machine-learning-databases/badges/badges.info) (it is stored in the same directory as 'badges.data') using which ever method you like (pandas, numpy, built-in python modules). *HINT: I recommend using pandas to read in this **fixed width** dataset.* 

The first column are the labels we want to predict (the + or - sign); the second column are the names. If you do this step correctly the data will be identical to the first 5 rows below:
```
[['+', 'Naoki Abe'],
 ['-', 'Myriam Abramson'],
 ['+', 'David W. Aha'],
 ['+', 'Kamal M. Ali'],
 ['-', 'Eric Allender']]
```

Depending on how you read in the data. Format your data so that it is in an acceptable format for [train_test_split](http://scikit-learn.org/stable/modules/generated/sklearn.model_selection.train_test_split.html) function in in sci-kit learn.

Next Split the data set the random_state parameter to be 0, and assign the approriate parameter to set aside 80% of your dataset for training data.

Now encode your labels using [LabelEncoder](http://scikit-learn.org/stable/modules/generated/sklearn.preprocessing.LabelEncoder.html#sklearn.preprocessing.LabelEncoder). 

Next create a [TfidfVectorizer](http://scikit-learn.org/stable/modules/generated/sklearn.feature_extraction.text.TfidfVectorizer.html#sklearn.feature_extraction.text.TfidfVectorizer) model using sci-kit learn's implementation using the default values.

For your training features learn the vocabulary and idf and return the term-document matrix. For your testing set **only** transform documents to a document term matrix.

Assign your encoded labels to train_y and test_y for the training and testing labels respectively.
Assign your document-term matrices to the training and testing document-term matrix.


<!-- grade_id: cell-22b77393ce3fb980, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-2b0f7b5ae4588ecb, points: 1, grade: true, solution: false -->
```python
assert_equal(train_y.tolist(), [0, 1, 0, 0, 1, 1, 1, 0, 1, 0, 0, 0, 0, 0, 1, 0, 1, 0, 0, 0, 1, 0, 0,
       0, 1, 1, 1, 0, 1, 1, 1, 1, 0, 0, 0, 0, 0, 0, 0, 1, 1, 1, 0, 0, 1, 0,
       1, 0, 1, 0, 1, 1, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 1, 0, 1, 0,
       0, 1, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 1, 1, 0, 1, 0, 0, 0, 0, 0, 1, 0,
       0, 0, 0, 0, 0, 0, 0, 0, 1, 1, 0, 0, 0, 0, 1, 0, 0, 1, 0, 1, 1, 0, 1,
       0, 1, 1, 0, 1, 0, 0, 0, 0, 0, 0, 1, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 1,
       0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 1, 1, 1, 0, 1, 1, 0, 0, 0, 1, 0, 1, 0,
       1, 0, 0, 1, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 1, 0, 0, 1, 0, 0, 0, 0, 0,
       0, 1, 0, 1, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0,
       1, 1, 0, 0, 0, 1, 0, 1, 1, 0, 1, 0, 1, 0, 1, 0, 1, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 1, 0])

assert_equal(test_y.tolist(), [1, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 1, 0, 0,
       1, 0, 1, 0, 0, 0, 0, 0, 1, 1, 1, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 1, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0])

```

# Problem 4: Fit a Naive Bayes Multinomial Model to the Badges dataset
Use sci-kit learn's implementation of a [naive bayes classifer for multinomial model](http://scikit-learn.org/stable/modules/generated/sklearn.naive_bayes.MultinomialNB.html) to fit a model to the Badges dataset. Lastly get predictions of the testing labels and name your predictions *_pred*

<!-- grade_id: cell-48de75d9feeddc8d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

```python
print(metrics.classification_report(_pred, test_y))
```

<!-- grade_id: cell-c075ab26e7d7a622, points: 1, grade: true, solution: false -->
```python
assert_equal(_pred.tolist(), [0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0])

```

# Problem 5: Using N-Grams to classify badges 
Use sci-kit learn's  [Count Vectorizer](http://scikit-learn.org/stable/modules/generated/sklearn.feature_extraction.text.CountVectorizer.html). Learn the names and transform the training term-document matrix.
For the testing names transform into the testing term-document matrix.

Next create another Multinomial Naive Bayes model.
Fit the model on the  training term-document matrix and training labels.
Next make predictions on the testing labels with the testing term-document matrix. store this value as *p*.


<!-- grade_id: cell-1015cdc20b7f9990, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

```python
print(metrics.classification_report(test_y, p))
```

<!-- grade_id: cell-4589c856766ce624, points: 1, grade: true, solution: false -->
```python
assert_equal(p.tolist(), [0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0])
```

As before, we will use the NLTK Reuters corpus which contains 10,788 documents which have been classified into 90 topics. The documents have file ids (held in the $\texttt{fileids()}$ method) which denotes if they are training or testing documents. An example is below.

```python
#Let's look at the file id for the first document
reuters.fileids()[0]
```

# Problem 6: Write a function that splits the reuters data into testing and training sets

Write a function called $\texttt{splitter}$ which takes in a corpora, and returns a list of file ids for the training documents, and a list of file ids for the testing documents.

<!-- grade_id: cell-1223e8a718a4bbde, points: 0, grade: false, solution: true -->
```python
def splitter(corpora):
    """

    Inputs
    ----------

    corpora: an nltk corpora
    
    Returns
    -------

    trainig_ids: a list of training document file ids
    testing_ids: a list of testing document file ids
    """

    # YOUR CODE HERE
    
    return training_ids, testing_ids
```

<!-- grade_id: cell-1a23052376286497, points: 1, grade: true, solution: false -->
```python
training_ids, testing_ids = splitter(reuters)
assert_equal(len(training_ids),7769)
assert_equal(len(testing_ids),3019)
```


We will use the below function, along with the training and testing ids that we found to get the full testing and training datasets


```python
def get_categories_from_fileids(corpus, fileids):
    """
    Finds categories for each element of 'fileids'.
    
    Parameters
    ----------
    corpus: An NLTK corpus.
    fileids: A list of strings.
    
    Returns
    -------
    A list of strings.
    """
    
    result = [sorted(corpus.categories(fileids=f))[0] for f in fileids]
    
    return result
```

```python
X_train = [reuters.raw(fileids=fileid) for fileid in training_ids][:1000]
y_train = get_categories_from_fileids(reuters, training_ids)[:1000]


X_test = [reuters.raw(fileids=fileid) for fileid in testing_ids]
y_test = get_categories_from_fileids(reuters, testing_ids)
```

# Problem 7: Fit a naive bayes classifier using different n-grams

Write a function called $\texttt{naivebayes_categories}$ which takes in a training set of data, a lower n-gram range, and an upper n-gram range. Furthermore, use english stop words in the CountVectorizer.

<!-- grade_id: cell-30b8f9c73b7265d9, points: 0, grade: false, solution: true -->
```python
def naivebayes_categories(X_train, y_train, lower, upper):
    """

    Inputs
    ----------
    X_train: The training data
    y_train: The training labels
    lower: the lower n_gram range
    upper: the upper n_gram range
    
    Returns
    -------
    clf: the fitted model
    
    """
    # YOUR CODE HERE
    
    return clf
```

<!-- grade_id: cell-d36114816765a759, points: 1, grade: true, solution: false -->
```python
model = naivebayes_categories(X_train,y_train,1,2)
predictions = model.predict(X_test)
assert_is_instance(model,Pipeline)
assert_equal(model.classes_[0],'acq')
assert_equal(predictions[0],'trade')
```

Let's see how accurate the model is.

```python
accuracy_score(predictions, y_test)
```

```python

```