# Module 11 Assignment

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

```python
import numpy as np
import pandas as pd
import datetime
import matplotlib.pyplot as plt
import seaborn as sns
%matplotlib inline

import warnings
warnings.filterwarnings('ignore')

from nose.tools import assert_equal, assert_true
import helper

from datetime import datetime as dt
from datetime import time
from datetime import date 
from dateutil import tz
from datetime import timedelta
```

```python
#get the current day
current_day = datetime.date.today()
```

# Problem 1: Write a function to get the time and date in a given city

Write a function called $\texttt{get_time_and_date}$ which, using tz from the dateutil library, takes in a string containing a country and city and returns the current time in that city in iso format. 

<!-- grade_id: cell-693337c2e9173d00, points: 0, grade: false, solution: true -->
```python
def get_time_and_date(timezone):
    '''
    Inputs
    -------
    
    Timezone: a string, the country/city to check the time and date in
    
    
    Returns
    --------
    
    date_and_time: a date and time in iso format
    
    
    '''
    
    # YOUR CODE HERE
    
    return date_and_time
```

<!-- grade_id: cell-f4331217723823e8, points: 1, grade: true, solution: false -->
```python
date_and_time = get_time_and_date('Spain/Barcelona')
assert_equal(date_and_time[:4],'2017')
assert_equal(len(date_and_time) > 10, 1)
```

# Problem 2: Write a function to get the past time

Write a function called $\texttt{get_past_time}$ which takes in a number of hours, and then returns the time that is $\texttt{hours}$ before the current time.

<!-- grade_id: cell-32d55078e5c4ddb4, points: 0, grade: false, solution: true -->
```python
def get_past_time(hours):
    '''
    Inputs
    -------
    
    hours: the number of hours to go back
    
    
    Returns
    --------
    
    past_time: the time in the past in iso format
    
    
    '''
    
    # YOUR CODE HERE
    
    return past_time
```

One note: Since we're dealing with times here, it's possible that you may not pass the first assert statement, especially if you're doing your homework late at night. It may be necessary to restart and run all cells in order to pass the assert statement.

<!-- grade_id: cell-dbbe96a68ed2be68, points: 1, grade: true, solution: false -->
```python
past_time = get_past_time(hours=24)
assert_equal(past_time[:4],'2017')
assert_equal(str(current_day.day-1) in past_time,1)
```

# Problem 3 : Preprocessing Stock Data
For this problem read in the dow jones dataset (which is in the same directory) using whichever method you prefer. Then based on how you read in data select the following columns: stock, date, open, high, low, close, percent_change_price and create a dataframe. Set the index to be the date column from the dow jones dataset (*Do not convert your index to datetime*). 

Save this panadas dataframe as df.


<!-- grade_id: cell-8bbb4993abc54c8f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

```python
df.head(5)
```

<!-- grade_id: cell-0cc630de094f497d, points: 1, grade: true, solution: false -->
```python
assert_true(type(df) == pd.core.frame.DataFrame, "You did not create a dataframe called df")
assert_equal(sorted(df.columns.tolist()), ['close', 'high', 'low', 'open', 'percent_change_price', 'stock'])
assert_equal(helper.df, df[['open', 'high', 'low', 'close']].values.tolist())
assert_equal(helper.index, df.index.tolist())
```

# Problem 4: Converting Index to Datetime Index

For this problem convert the index to a datetime index.

<!-- grade_id: cell-c9a4a36949a16cb7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-1c758bac8933c7e2, points: 1, grade: true, solution: false -->
```python
assert_true(type(df.index) == pd.tseries.index.DatetimeIndex, msg = 'You did not create a datetime index')
assert_equal(df.index.values.tolist()[61:205], helper.index_65_205)
```

```python
# Let's take a look at our data
g = sns.pairplot(df, hue='stock', size=3, markers="+", diag_kind="kde")
```

In the cell below we'll take a look at percent_change_price for MSFT, IBM, and CSCO.

```python
df[(df['stock'] == 'IBM') | (df['stock'] == 'MSFT') | (df['stock'] == 'CSCO')].percent_change_price.plot(figsize=(12, 6))
```

# Problem 5: Fourier Transforms 

For this problem compute a one-dimensional discrete Fourier Transform on *data* save this variable as yfft. Then compute the Discrete Fourier Transform sample frequencies (set the Sample spacing  to .1) and save it in a variable called freq_yfft.


```python
data = df[(df['stock'] == 'IBM') | (df['stock'] == 'MSFT') | (df['stock'] == 'CSCO')].percent_change_price
```

<!-- grade_id: cell-d5dbeb683ac5d982, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-a4f925f786a95929, points: 1, grade: true, solution: false -->
```python
assert_equal(yfft.tolist(), helper.yf)
assert_equal(freq_yfft.tolist(), helper.freq)
```

# Problem 6: Visualizing Results

For this problem, we will plot the signal from problem 5. Plot it using matplotlib with a fig and ax object. Title the plot "signal" and label the x-axis as "Frequency". Finally, set the limits of the x-axis to be 0 to 4.

<!-- grade_id: cell-88678889904a65e4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-45bfef3ce30a2fae, points: 1, grade: true, solution: false -->
```python
assert_equal(ax.get_xlabel(),'Frequency')
assert_equal(ax.get_title(),'Signal')
assert_equal(ax.get_xlim()[0],0.0)
assert_equal(ax.get_xlim()[1],4.0)
assert_equal(ax.lines[0].get_data()[0][3],0.4)
assert_equal(len(ax.lines[0].get_data()[0]),75)
```