# Module 4 Assignment

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

```python
import numpy as np
import pandas as pd
from sklearn.neural_network import MLPClassifier, MLPRegressor
from sklearn.model_selection import train_test_split
from sklearn.preprocessing import LabelEncoder
import pandas as pd
import numpy as np
import random
from nose.tools import assert_equal, assert_almost_equal
from keras.utils import np_utils
from keras.layers.core import Dropout
from keras.models import Sequential
from keras.layers.core import Dense, Activation
from keras.optimizers import RMSprop
from keras.utils.np_utils import to_categorical
from nose.tools import assert_is_instance, assert_equal
from numpy.testing import assert_array_equal
```

Let's consider a hypothetical dataset containing movie watchers and their reviews on a scale from 0 to 5 of 7 different movies.

```python
user_ids = ["Robert", "Edward", "Kelechi", "Will"]
movie_ids = list(range(7))

data = np.array(
    [[4, 0, 1, 5, 1, 2, 3],
     [5, 5, 4, 3, 1, 0, 1],
     [2, 3, 4, 1, 0, 5, 1],
     [1, 2, 2, 1, 0, 1, 4]]
)

movies_df = pd.DataFrame(
    data=data,
    index=user_ids,
    columns=movie_ids
)

print(movies_df)
```

# Problem 1: Create the matrix of favorable ratings

Write a function called "$\texttt{favorable_matrix}$" which takes in a dataframe of movie ratings and changes any rating above or equal to "$\texttt{min_rating}$" to a 1 and any rating below "$\texttt{min_rating}$" to a 0, returning a numpy matrix.

<!-- grade_id: cell-3458304c047b9b61, points: 0, grade: false, solution: true -->
```python
def favorable_matrix(df,min_rating):
    """
    Takes a pandas data frame and returns a numpy matrix
    with 1 for ratings > 3 and 0 for ratings <= 3.
    
    Parameters
    ----------
    df: A pandas.DataFrame
    min_rating: the minimum rating for the review to be favorable
    
    Returns
    -------
    rating_matrix: A 2-d numpy array of favorable ratings
    """
    
    # YOUR CODE HERE
    
    return rating_matrix
```

<!-- grade_id: cell-a62da77714c763bf, points: 1, grade: true, solution: false -->
```python
rating_matrix = favorable_matrix(movies_df,3)

assert_is_instance(rating_matrix,np.ndarray)
assert_array_equal(favorable_matrix(movies_df,3)[1],np.array([1, 1, 1, 1, 0, 0, 0]))
```

Now let's define the cosine similarity function so that we can come up with recommendations.

```python
# Define the Cosine Similarity function

def cosine_similarity(u, v):
    return(np.dot(u, v)/np.sqrt((np.dot(u, u) * np.dot(v, v))))
```

# Problem 2: Find similar users

Write a function called "$\texttt{find_similar}$" which takes in a given user's movie scores and compares them to the scores in a matrix of other user's ratings. The function should then output which user in our dataset is most similar to the candidate user.

<!-- grade_id: cell-8f3fb117aca2827e, points: 0, grade: false, solution: true -->
```python
def find_similar(ratings_matrix, user_ratings):
    """
    Parameters
    ----------
    ratings_matrix: the matrix of past user ratings
    user_ratings: the ratings of the given user
    
    Returns
    -------
    most_similar_index: the row of the matrix that is most similar to the given user
    """
    
    # YOUR CODE HERE
    
    return most_similar_index
```

<!-- grade_id: cell-f9dc5d68673ebd31, points: 1, grade: true, solution: false -->
```python
user_ratings = np.array([0,1,0,1,0,0,1])
most_similar_index = find_similar(rating_matrix, user_ratings)
assert_equal(movies_df.index[most_similar_index],'Robert')
```

<!-- grade_id: cell-39bbef9149a90a02, points: 1, grade: true, solution: false -->
```python
user_ratings = np.array([0,1,1,1,0,0,1])
most_similar_index = find_similar(rating_matrix, user_ratings)
assert_equal(movies_df.index[most_similar_index],'Edward')
```

Now let's switch gears and load in the familiar breast cancer data

```python
df = pd.read_csv('./breast-cancer-wisconsin.data')
df.drop('id', axis=1, inplace=True)
df.head()
```

# Problem 3: Splitting Data into Training and Testing sets
For this problem use [train_test_split](http://scikit-learn.org/stable/modules/generated/sklearn.model_selection.train_test_split.html) function from sci-kit learn on the dataframe that contains the breast cancer dataset. Set the random_state argument to be 0 and use 20% of the data for testing. The labels for this dataset are the classes and the features are the other columns.

Convert the dataframe and series objects to numpy arrays.

After you split the data encode the Labels using [LabelEncoder](http://scikit-learn.org/stable/modules/generated/sklearn.preprocessing.LabelEncoder.html).  

The name of encoded Labels for the training and testing set should be: *train_y* & *test_y*.  
The name of the Features for the training and testing set should be: *train_X* & *test_X*.

```python
random.seed(0)
```

<!-- grade_id: cell-03f92af120e221da, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-6050a3efc0aa700a, points: 1, grade: true, solution: false -->
```python
# Checking Type
for x in [train_y, train_X, test_y, test_X]:
    assert_equal(type(x), np.ndarray)

# Checking Shape 
assert_equal(train_X.shape, (546, 9))
assert_equal(train_y.shape, (546, ))
assert_equal(test_y.shape, (137, ))
assert_equal(test_X.shape, (137, 9))

# Checking Data
assert_equal(train_X.tolist(), [[10, 1, 1, 1, 2, 10, 5, 4, 1], [1, 1, 1, 1, 1, 1, 3, 1, 1], [5, 1, 1, 1, 2, 1, 3, 1, 1], [3, 1, 2, 1, 2, 1, 2, 1, 1], [8, 2, 3, 1, 6, 3, 7, 1, 1], [1, 1, 1, 1, 2, 1, 3, 1, 1], [2, 1, 1, 1, 2, 1, 3, 1, 1], [4, 1, 1, 3, 1, 1, 2, 1, 1], [5, 7, 7, 1, 5, 8, 3, 4, 1], [8, 10, 10, 8, 6, 9, 3, 10, 10], [1, 1, 1, 1, 2, 1, 2, 1, 1], [6, 1, 3, 1, 4, 5, 5, 10, 1], [3, 1, 1, 1, 2, 1, 2, 1, 1], [5, 10, 10, 5, 4, 5, 4, 4, 1], [5, 3, 3, 1, 3, 3, 3, 3, 3], [3, 1, 1, 1, 3, 1, 2, 1, 1], [2, 1, 1, 1, 2, 1, 2, 1, 1], [1, 1, 1, 1, 1, 1, 1, 1, 1], [4, 2, 1, 1, 2, 1, 1, 1, 1], [3, 10, 8, 7, 6, 9, 9, 3, 8], [10, 7, 7, 4, 5, 10, 5, 7, 2], [1, 1, 1, 1, 3, 1, 1, 1, 1], [5, 10, 10, 3, 7, 3, 8, 10, 2], [3, 10, 7, 8, 5, 8, 7, 4, 1], [10, 4, 4, 6, 2, 10, 2, 3, 1], [8, 7, 8, 5, 5, 10, 9, 10, 1], [5, 3, 3, 3, 2, 3, 4, 4, 1], [10, 8, 10, 10, 6, 1, 3, 1, 10], [4, 1, 4, 1, 2, 1, 1, 1, 1], [5, 1, 2, 1, 2, 1, 1, 1, 1], [4, 1, 1, 1, 2, 2, 3, 2, 1], [5, 1, 1, 1, 2, 1, 3, 1, 2], [2, 2, 2, 1, 1, 1, 7, 1, 1], [2, 1, 1, 1, 2, 1, 3, 1, 1], [7, 6, 10, 5, 3, 10, 9, 10, 2], [8, 3, 8, 3, 4, 9, 8, 9, 8], [8, 10, 10, 10, 6, 10, 10, 10, 1], [1, 1, 1, 1, 2, 1, 1, 1, 1], [1, 1, 3, 1, 1, 1, 2, 1, 1], [1, 1, 1, 1, 2, 1, 1, 1, 1], [5, 6, 6, 2, 4, 10, 3, 6, 1], [3, 1, 1, 1, 2, 1, 2, 1, 1], [2, 1, 3, 2, 2, 1, 2, 1, 1], [4, 6, 6, 5, 7, 6, 7, 7, 3], [7, 4, 6, 4, 6, 1, 4, 3, 1], [5, 1, 3, 3, 2, 2, 2, 3, 1], [10, 6, 4, 3, 10, 10, 9, 10, 1], [1, 1, 1, 1, 2, 10, 3, 1, 1], [1, 1, 1, 1, 2, 1, 2, 1, 1], [9, 4, 5, 10, 6, 10, 4, 8, 1], [8, 10, 10, 10, 6, 10, 10, 10, 10], [1, 1, 1, 1, 2, 1, 3, 1, 1], [1, 1, 1, 1, 2, 1, 2, 1, 1], [3, 1, 1, 1, 2, 1, 3, 1, 1], [3, 1, 4, 1, 2, 1, 1, 1, 1], [3, 1, 1, 1, 2, 1, 3, 1, 1], [1, 1, 1, 1, 2, 1, 1, 1, 1], [3, 1, 1, 1, 2, 1, 3, 2, 1], [1, 1, 3, 1, 2, 1, 1, 1, 1], [4, 1, 1, 1, 2, 1, 3, 1, 1], [4, 1, 1, 2, 2, 1, 1, 1, 1], [4, 8, 8, 5, 4, 5, 10, 4, 1], [1, 1, 2, 2, 2, 1, 3, 1, 1], [5, 1, 1, 1, 2, 1, 3, 1, 1], [1, 1, 3, 1, 2, 1, 2, 1, 1], [3, 1, 2, 2, 2, 1, 1, 1, 1], [3, 1, 1, 1, 1, 1, 2, 1, 1], [10, 8, 4, 4, 4, 10, 3, 10, 4], [8, 6, 4, 3, 5, 9, 3, 1, 1], [4, 1, 1, 1, 1, 1, 2, 1, 1], [8, 10, 8, 8, 4, 8, 7, 7, 1], [5, 2, 2, 2, 2, 1, 2, 2, 1], [8, 4, 4, 5, 4, 7, 7, 8, 2], [4, 4, 4, 4, 6, 5, 7, 3, 1], [10, 10, 10, 7, 10, 10, 8, 2, 1], [3, 1, 2, 1, 2, 1, 2, 1, 1], [1, 1, 1, 1, 2, 1, 1, 1, 1], [1, 1, 1, 1, 1, 1, 3, 1, 1], [2, 1, 1, 1, 2, 1, 1, 1, 1], [6, 6, 7, 10, 3, 10, 8, 10, 2], [9, 7, 7, 5, 5, 10, 7, 8, 3], [1, 2, 1, 3, 2, 1, 2, 1, 1], [5, 2, 2, 4, 2, 4, 1, 1, 1], [9, 10, 10, 10, 10, 5, 10, 10, 10], [10, 10, 10, 6, 8, 4, 8, 5, 1], [2, 1, 1, 1, 2, 1, 2, 1, 1], [3, 1, 1, 1, 3, 1, 2, 1, 1], [5, 3, 3, 3, 6, 10, 3, 1, 1], [8, 7, 6, 4, 4, 10, 5, 1, 1], [1, 1, 1, 3, 2, 1, 1, 1, 1], [8, 10, 5, 3, 8, 4, 4, 10, 3], [4, 1, 1, 1, 2, 1, 3, 1, 1], [7, 5, 6, 10, 5, 10, 7, 9, 4], [10, 4, 3, 10, 4, 10, 10, 1, 1], [2, 1, 1, 1, 2, 1, 1, 1, 1], [2, 1, 1, 1, 2, 1, 2, 1, 1], [1, 1, 1, 1, 2, 1, 3, 1, 1], [9, 8, 8, 9, 6, 3, 4, 1, 1], [4, 1, 1, 1, 2, 3, 2, 1, 1], [9, 6, 9, 2, 10, 6, 2, 9, 10], [1, 2, 1, 3, 2, 1, 1, 2, 1], [5, 3, 2, 8, 5, 10, 8, 1, 2], [3, 1, 1, 1, 2, 1, 2, 1, 2], [1, 1, 1, 1, 2, 1, 1, 1, 1], [3, 1, 1, 2, 2, 1, 1, 1, 1], [1, 1, 1, 1, 1, 1, 1, 3, 1], [2, 1, 1, 1, 1, 1, 3, 1, 1], [10, 10, 10, 8, 2, 10, 4, 1, 1], [5, 1, 1, 1, 2, 1, 3, 1, 1], [4, 1, 1, 3, 2, 1, 3, 1, 1], [5, 10, 10, 3, 8, 1, 5, 10, 3], [1, 1, 1, 1, 2, 1, 1, 1, 1], [3, 1, 1, 1, 1, 1, 2, 1, 1], [10, 8, 8, 2, 8, 10, 4, 8, 10], [1, 1, 1, 1, 2, 1, 3, 1, 1], [4, 1, 2, 1, 2, 1, 2, 1, 1], [1, 1, 1, 1, 2, 1, 3, 1, 1], [2, 1, 1, 2, 2, 1, 3, 1, 1], [6, 1, 1, 1, 2, 1, 3, 1, 1], [6, 5, 4, 4, 3, 9, 7, 8, 3], [5, 1, 2, 10, 4, 5, 2, 1, 1], [4, 1, 1, 1, 2, 1, 1, 1, 1], [10, 10, 10, 3, 10, 10, 9, 10, 1], [5, 1, 1, 1, 2, 1, 1, 1, 1], [10, 10, 10, 10, 10, 1, 8, 8, 8], [1, 3, 1, 2, 2, 2, 5, 3, 2], [6, 1, 3, 1, 2, 1, 3, 1, 1], [8, 7, 8, 5, 10, 10, 7, 2, 1], [6, 10, 10, 10, 4, 10, 7, 10, 1], [9, 10, 10, 1, 10, 8, 3, 3, 1], [1, 1, 1, 1, 2, 1, 2, 1, 1], [1, 3, 1, 1, 2, 1, 2, 2, 1], [3, 1, 1, 1, 2, 1, 3, 1, 1], [1, 2, 3, 1, 2, 1, 1, 1, 1], [1, 1, 1, 3, 2, 1, 1, 1, 1], [2, 1, 1, 1, 2, 1, 3, 1, 1], [5, 10, 10, 10, 5, 2, 8, 5, 1], [7, 4, 7, 4, 3, 7, 7, 6, 1], [5, 1, 1, 1, 2, 1, 2, 1, 1], [4, 1, 1, 1, 2, 1, 1, 1, 1], [4, 1, 1, 1, 2, 1, 1, 1, 1], [6, 8, 8, 1, 3, 4, 3, 7, 1], [4, 1, 1, 1, 2, 1, 2, 1, 1], [5, 1, 1, 1, 2, 1, 2, 1, 1], [2, 3, 1, 1, 5, 1, 1, 1, 1], [10, 5, 10, 3, 5, 8, 7, 8, 3], [5, 2, 2, 2, 1, 1, 2, 1, 1], [3, 1, 1, 1, 2, 1, 1, 1, 1], [6, 8, 7, 5, 6, 8, 8, 9, 2], [1, 1, 1, 1, 1, 1, 2, 1, 1], [2, 1, 1, 1, 2, 1, 3, 1, 1], [1, 1, 1, 1, 1, 1, 1, 1, 1], [5, 1, 1, 1, 2, 1, 1, 1, 1], [5, 1, 1, 1, 2, 1, 1, 1, 1], [3, 1, 1, 1, 2, 1, 1, 1, 1], [4, 1, 1, 1, 2, 1, 1, 1, 1], [5, 10, 10, 10, 10, 2, 10, 10, 10], [4, 2, 2, 1, 2, 1, 2, 1, 1], [3, 3, 2, 1, 2, 3, 3, 1, 1], [6, 1, 1, 3, 2, 1, 1, 1, 1], [7, 3, 2, 10, 5, 10, 5, 4, 4], [5, 3, 4, 1, 8, 10, 4, 9, 1], [2, 1, 1, 1, 3, 1, 2, 1, 1], [8, 3, 5, 4, 5, 10, 1, 6, 2], [1, 1, 1, 1, 2, 1, 1, 1, 1], [1, 1, 1, 1, 2, 1, 1, 1, 1], [3, 1, 1, 1, 2, 1, 1, 1, 1], [4, 10, 8, 5, 4, 1, 10, 1, 1], [3, 1, 1, 1, 2, 1, 1, 1, 1], [3, 1, 1, 1, 2, 1, 2, 1, 1], [2, 1, 2, 1, 2, 1, 3, 1, 1], [4, 8, 6, 4, 3, 4, 10, 6, 1], [5, 4, 6, 10, 2, 10, 4, 1, 1], [3, 1, 3, 1, 3, 4, 1, 1, 1], [7, 5, 10, 10, 10, 10, 4, 10, 3], [1, 1, 1, 1, 2, 1, 1, 1, 1], [2, 1, 1, 1, 2, 1, 3, 1, 1], [8, 10, 10, 1, 3, 6, 3, 9, 1], [2, 1, 1, 1, 2, 1, 3, 1, 1], [2, 1, 1, 1, 2, 1, 1, 1, 1], [4, 1, 2, 1, 2, 1, 3, 1, 1], [5, 1, 1, 1, 2, 1, 2, 1, 1], [5, 1, 1, 3, 2, 1, 1, 1, 1], [5, 2, 2, 2, 2, 1, 1, 1, 2], [1, 1, 1, 1, 2, 1, 3, 1, 1], [5, 2, 3, 1, 6, 10, 5, 1, 1], [9, 10, 10, 1, 10, 8, 3, 3, 1], [8, 7, 8, 2, 4, 2, 5, 10, 1], [8, 4, 7, 1, 3, 10, 3, 9, 2], [2, 3, 1, 1, 3, 1, 1, 1, 1], [3, 3, 5, 2, 3, 10, 7, 1, 1], [5, 1, 1, 1, 3, 2, 2, 2, 1], [10, 10, 10, 10, 3, 10, 10, 6, 1], [3, 1, 1, 1, 2, 4, 1, 1, 1], [10, 6, 3, 6, 4, 10, 7, 8, 4], [6, 10, 10, 10, 8, 10, 10, 10, 7], [2, 5, 3, 3, 6, 7, 7, 5, 1], [1, 1, 1, 1, 1, 1, 2, 1, 1], [5, 1, 1, 1, 2, 1, 2, 1, 1], [8, 6, 4, 10, 10, 1, 3, 5, 1], [5, 1, 1, 3, 4, 1, 3, 2, 1], [1, 1, 1, 1, 1, 1, 1, 1, 1], [3, 1, 1, 1, 2, 1, 2, 1, 1], [1, 1, 1, 2, 1, 1, 1, 1, 1], [7, 5, 3, 7, 4, 10, 7, 5, 5], [9, 9, 10, 3, 6, 10, 7, 10, 6], [5, 1, 1, 3, 2, 1, 1, 1, 1], [1, 1, 1, 1, 2, 2, 2, 1, 1], [5, 1, 1, 1, 2, 1, 2, 2, 1], [6, 8, 7, 8, 6, 8, 8, 9, 1], [1, 1, 1, 1, 2, 1, 1, 1, 1], [4, 4, 2, 1, 2, 5, 2, 1, 2], [8, 10, 10, 7, 10, 10, 7, 3, 8], [4, 3, 2, 1, 3, 1, 2, 1, 1], [3, 2, 2, 1, 2, 1, 2, 3, 1], [3, 1, 1, 1, 2, 1, 3, 1, 1], [1, 1, 1, 2, 1, 3, 1, 1, 7], [5, 5, 5, 8, 10, 8, 7, 3, 7], [3, 6, 4, 10, 3, 3, 3, 4, 1], [5, 1, 1, 1, 2, 1, 1, 1, 1], [2, 1, 1, 1, 2, 1, 2, 2, 1], [3, 3, 2, 6, 3, 3, 3, 5, 1], [1, 1, 1, 1, 2, 1, 2, 1, 1], [3, 1, 1, 1, 2, 1, 3, 1, 1], [1, 1, 4, 1, 2, 1, 2, 1, 1], [1, 1, 1, 1, 2, 1, 3, 1, 1], [2, 1, 1, 1, 2, 1, 1, 1, 1], [4, 1, 1, 3, 2, 1, 3, 1, 1], [1, 1, 1, 1, 2, 1, 1, 1, 1], [5, 1, 1, 1, 2, 1, 2, 1, 1], [8, 5, 6, 2, 3, 10, 6, 6, 1], [1, 1, 1, 1, 2, 1, 1, 1, 1], [8, 10, 10, 8, 7, 10, 9, 7, 1], [2, 1, 1, 1, 2, 1, 2, 1, 1], [4, 1, 1, 2, 2, 1, 2, 1, 1], [3, 1, 1, 1, 2, 1, 1, 1, 1], [1, 1, 1, 1, 2, 1, 3, 1, 1], [5, 1, 1, 1, 2, 1, 2, 1, 1], [3, 1, 1, 1, 2, 1, 2, 1, 1], [10, 10, 7, 8, 7, 1, 10, 10, 3], [8, 10, 10, 10, 8, 10, 10, 7, 3], [5, 5, 7, 8, 6, 10, 7, 4, 1], [10, 5, 7, 3, 3, 7, 3, 3, 8], [2, 1, 1, 1, 2, 1, 3, 1, 1], [1, 1, 1, 1, 2, 1, 1, 1, 1], [10, 5, 7, 4, 4, 10, 8, 9, 1], [4, 2, 4, 3, 2, 2, 2, 1, 1], [8, 6, 7, 3, 3, 10, 3, 4, 2], [6, 1, 1, 1, 1, 1, 1, 1, 1], [10, 10, 10, 2, 10, 10, 5, 3, 3], [3, 1, 1, 1, 2, 2, 7, 1, 1], [3, 1, 1, 1, 2, 1, 2, 3, 1], [1, 1, 1, 1, 2, 3, 3, 1, 1], [1, 1, 1, 1, 2, 4, 1, 1, 1], [10, 10, 10, 10, 5, 10, 10, 10, 7], [1, 1, 1, 1, 2, 1, 2, 1, 1], [5, 6, 7, 8, 8, 10, 3, 10, 3], [1, 1, 1, 1, 2, 1, 1, 1, 1], [5, 1, 1, 2, 2, 2, 3, 1, 1], [1, 1, 1, 1, 1, 1, 1, 1, 1], [6, 9, 7, 5, 5, 8, 4, 2, 1], [6, 3, 2, 1, 3, 4, 4, 1, 1], [5, 3, 2, 4, 2, 1, 1, 1, 1], [2, 3, 1, 1, 2, 1, 2, 1, 1], [10, 8, 8, 4, 10, 10, 8, 1, 1], [5, 1, 1, 4, 2, 1, 3, 1, 1], [1, 1, 1, 1, 4, 3, 1, 1, 1], [4, 1, 1, 1, 2, 1, 2, 1, 1], [1, 1, 1, 1, 2, 1, 2, 1, 1], [3, 1, 1, 1, 2, 1, 1, 1, 1], [1, 1, 1, 1, 2, 1, 2, 3, 1], [4, 1, 1, 1, 2, 1, 2, 1, 1], [10, 4, 2, 1, 3, 2, 4, 3, 10], [3, 2, 1, 1, 2, 1, 2, 2, 1], [1, 6, 8, 10, 8, 10, 5, 7, 1], [5, 2, 1, 1, 2, 1, 1, 1, 1], [1, 1, 3, 1, 2, 1, 2, 1, 1], [1, 1, 1, 1, 2, 1, 1, 1, 1], [10, 3, 5, 4, 3, 7, 3, 5, 3], [5, 1, 1, 1, 2, 1, 3, 2, 1], [3, 2, 2, 3, 2, 3, 3, 1, 1], [4, 1, 1, 1, 2, 1, 1, 1, 1], [1, 1, 1, 1, 2, 1, 3, 1, 1], [1, 1, 1, 1, 1, 1, 2, 1, 1], [5, 10, 6, 1, 10, 4, 4, 10, 10], [5, 1, 1, 1, 2, 1, 2, 1, 1], [3, 2, 2, 1, 4, 3, 2, 1, 1], [2, 1, 1, 1, 1, 1, 1, 1, 1], [1, 1, 1, 1, 2, 1, 3, 1, 1], [8, 8, 7, 4, 10, 10, 7, 8, 7], [2, 1, 1, 1, 2, 1, 2, 1, 1], [1, 1, 1, 1, 2, 1, 1, 1, 1], [10, 4, 3, 1, 3, 3, 6, 5, 2], [4, 1, 1, 1, 1, 1, 2, 1, 1], [1, 1, 1, 1, 2, 1, 1, 1, 1], [3, 1, 1, 1, 2, 1, 2, 1, 1], [4, 1, 1, 1, 2, 1, 3, 2, 1], [6, 5, 5, 8, 4, 10, 3, 4, 1], [5, 10, 10, 6, 10, 10, 10, 6, 5], [6, 2, 3, 1, 2, 1, 1, 1, 1], [4, 1, 1, 1, 2, 1, 2, 1, 1], [4, 1, 1, 1, 2, 1, 1, 1, 1], [1, 1, 3, 1, 2, 1, 1, 1, 1], [6, 10, 2, 8, 10, 2, 7, 8, 10], [3, 1, 1, 1, 2, 2, 3, 1, 1], [1, 1, 1, 1, 2, 1, 2, 1, 1], [8, 9, 9, 5, 3, 5, 7, 7, 1], [5, 1, 1, 1, 2, 1, 2, 1, 1], [1, 1, 1, 1, 1, 1, 3, 1, 1], [1, 1, 1, 1, 2, 1, 3, 1, 1], [4, 1, 1, 1, 2, 1, 3, 6, 1], [7, 5, 6, 10, 4, 10, 5, 3, 1], [5, 7, 10, 6, 5, 10, 7, 5, 1], [8, 3, 3, 1, 2, 2, 3, 2, 1], [7, 8, 8, 7, 3, 10, 7, 2, 3], [2, 1, 1, 1, 2, 1, 1, 1, 1], [8, 4, 4, 1, 6, 10, 2, 5, 2], [5, 1, 1, 1, 2, 1, 3, 1, 1], [6, 1, 1, 1, 2, 1, 3, 1, 1], [3, 1, 1, 1, 2, 1, 2, 1, 1], [7, 6, 6, 3, 2, 10, 7, 1, 1], [8, 6, 5, 4, 3, 10, 6, 1, 1], [3, 10, 3, 10, 6, 10, 5, 1, 4], [1, 1, 1, 2, 2, 1, 2, 1, 1], [1, 1, 3, 2, 2, 1, 3, 1, 1], [3, 1, 1, 1, 2, 1, 2, 1, 1], [4, 1, 1, 1, 2, 1, 2, 1, 1], [3, 1, 1, 1, 2, 1, 1, 1, 1], [4, 1, 1, 1, 2, 1, 3, 1, 1], [4, 1, 2, 1, 2, 1, 2, 1, 1], [3, 1, 1, 1, 2, 5, 1, 1, 1], [10, 3, 4, 5, 3, 10, 4, 1, 1], [8, 8, 9, 4, 5, 10, 7, 8, 1], [3, 1, 1, 1, 2, 1, 2, 1, 1], [3, 1, 1, 1, 2, 1, 3, 1, 1], [1, 1, 1, 1, 2, 1, 1, 1, 1], [4, 1, 1, 1, 2, 1, 2, 1, 1], [3, 1, 1, 3, 2, 1, 2, 1, 1], [6, 3, 3, 3, 3, 2, 6, 1, 1], [5, 1, 3, 1, 2, 1, 3, 1, 1], [10, 4, 4, 10, 6, 10, 5, 5, 1], [10, 4, 3, 2, 3, 10, 5, 3, 2], [5, 1, 1, 1, 2, 2, 3, 3, 1], [3, 1, 2, 1, 2, 1, 2, 1, 1], [1, 1, 1, 1, 2, 1, 3, 1, 1], [3, 2, 2, 3, 2, 1, 1, 1, 1], [10, 10, 10, 7, 9, 10, 7, 10, 10], [3, 7, 7, 4, 4, 9, 4, 8, 1], [5, 8, 8, 10, 5, 10, 8, 10, 3], [4, 1, 1, 1, 3, 1, 2, 2, 1], [6, 3, 4, 1, 5, 2, 3, 9, 1], [5, 3, 2, 1, 3, 1, 1, 1, 1], [5, 10, 10, 9, 6, 10, 7, 10, 5], [5, 1, 1, 6, 3, 1, 2, 1, 1], [2, 1, 1, 1, 2, 1, 2, 1, 1], [2, 3, 2, 2, 2, 2, 3, 1, 1], [3, 2, 1, 2, 2, 1, 3, 1, 1], [1, 1, 1, 1, 2, 1, 3, 1, 1], [3, 3, 5, 2, 3, 10, 7, 1, 1], [3, 1, 1, 4, 3, 1, 2, 2, 1], [5, 7, 9, 8, 6, 10, 8, 10, 1], [10, 5, 5, 6, 3, 10, 7, 9, 2], [1, 1, 1, 1, 1, 1, 2, 1, 1], [5, 1, 1, 1, 2, 1, 2, 1, 1], [4, 3, 3, 1, 2, 1, 3, 3, 1], [3, 1, 1, 1, 2, 5, 5, 1, 1], [1, 1, 2, 1, 2, 2, 4, 2, 1], [1, 1, 1, 1, 2, 1, 2, 1, 1], [1, 1, 1, 1, 2, 1, 3, 1, 1], [10, 10, 10, 8, 6, 1, 8, 9, 1], [5, 1, 2, 1, 2, 1, 3, 1, 1], [3, 1, 1, 1, 1, 1, 2, 1, 1], [3, 1, 1, 3, 2, 1, 1, 1, 1], [5, 4, 6, 8, 4, 1, 8, 10, 1], [5, 1, 1, 1, 2, 1, 2, 2, 1], [5, 1, 2, 1, 2, 1, 1, 1, 1], [7, 5, 6, 3, 3, 8, 7, 4, 1], [6, 10, 5, 5, 4, 10, 6, 10, 1], [10, 10, 10, 1, 6, 1, 2, 8, 1], [2, 1, 1, 1, 2, 1, 1, 1, 1], [6, 10, 7, 7, 6, 4, 8, 10, 2], [4, 8, 6, 3, 4, 10, 7, 1, 1], [2, 5, 7, 6, 4, 10, 7, 6, 1], [1, 2, 2, 1, 2, 1, 1, 1, 1], [3, 1, 1, 3, 1, 1, 3, 1, 1], [1, 1, 1, 1, 2, 1, 3, 2, 1], [5, 2, 3, 4, 2, 7, 3, 6, 1], [3, 1, 1, 1, 2, 1, 2, 1, 1], [4, 1, 1, 3, 1, 5, 2, 1, 1], [2, 1, 1, 1, 2, 1, 3, 1, 1], [4, 4, 4, 2, 2, 3, 2, 1, 1], [10, 10, 10, 3, 10, 8, 8, 1, 1], [1, 1, 1, 1, 2, 1, 2, 1, 1], [10, 10, 10, 10, 7, 10, 7, 10, 4], [4, 5, 5, 8, 6, 10, 10, 7, 1], [10, 4, 6, 1, 2, 10, 5, 3, 1], [5, 2, 4, 1, 1, 1, 1, 1, 1], [3, 1, 1, 1, 2, 1, 2, 1, 1], [5, 1, 1, 1, 2, 2, 2, 1, 1], [9, 1, 2, 6, 4, 10, 7, 7, 2], [5, 1, 2, 1, 2, 1, 1, 1, 1], [10, 3, 3, 1, 2, 10, 7, 6, 1], [5, 3, 5, 1, 8, 10, 5, 3, 1], [10, 4, 5, 4, 3, 5, 7, 3, 1], [3, 3, 2, 1, 3, 1, 3, 6, 1], [5, 1, 1, 3, 2, 1, 1, 1, 1], [3, 2, 1, 1, 2, 2, 3, 1, 1], [4, 8, 7, 10, 4, 10, 7, 5, 1], [3, 2, 1, 1, 1, 1, 2, 1, 1], [10, 3, 3, 10, 2, 10, 7, 3, 3], [4, 1, 1, 1, 2, 3, 1, 1, 1], [1, 2, 2, 1, 2, 1, 2, 1, 1], [4, 1, 1, 1, 2, 1, 2, 1, 1], [4, 1, 2, 1, 2, 1, 1, 1, 1], [9, 5, 5, 4, 4, 5, 4, 3, 3], [10, 8, 8, 2, 3, 4, 8, 7, 8], [1, 1, 1, 1, 2, 1, 3, 1, 1], [4, 1, 1, 1, 2, 1, 2, 1, 1], [3, 1, 1, 1, 2, 1, 1, 1, 1], [10, 9, 8, 7, 6, 4, 7, 10, 3], [9, 10, 10, 10, 10, 10, 10, 10, 1], [6, 1, 1, 1, 2, 1, 3, 1, 1], [2, 1, 1, 1, 2, 1, 1, 1, 1], [7, 8, 7, 6, 4, 3, 8, 8, 4], [1, 1, 1, 1, 2, 1, 3, 1, 1], [8, 4, 6, 3, 3, 1, 4, 3, 1], [2, 1, 1, 2, 2, 1, 1, 1, 1], [7, 4, 4, 3, 4, 10, 6, 9, 1], [5, 6, 6, 8, 6, 10, 4, 10, 4], [8, 2, 1, 1, 5, 1, 1, 1, 1], [5, 4, 5, 1, 8, 1, 3, 6, 1], [5, 1, 3, 1, 2, 1, 1, 1, 1], [10, 10, 10, 8, 6, 8, 7, 10, 1], [3, 3, 2, 2, 3, 1, 1, 2, 3], [4, 1, 1, 1, 2, 1, 2, 1, 1], [1, 1, 1, 1, 2, 1, 1, 1, 1], [7, 4, 5, 10, 2, 10, 3, 8, 2], [1, 1, 1, 1, 2, 1, 2, 1, 1], [5, 8, 4, 10, 5, 8, 9, 10, 1], [2, 1, 1, 2, 2, 1, 3, 1, 1], [5, 2, 1, 1, 2, 1, 3, 1, 1], [7, 6, 3, 2, 5, 10, 7, 4, 6], [2, 1, 1, 1, 2, 1, 2, 1, 1], [1, 1, 1, 1, 2, 1, 2, 1, 2], [4, 1, 1, 1, 2, 1, 3, 1, 1], [6, 2, 1, 1, 1, 1, 7, 1, 1], [5, 3, 5, 5, 3, 3, 4, 10, 1], [5, 6, 5, 6, 10, 1, 3, 1, 1], [5, 3, 6, 1, 2, 1, 1, 1, 1], [1, 1, 1, 1, 2, 1, 2, 1, 1], [8, 7, 8, 7, 5, 5, 5, 10, 2], [10, 10, 8, 6, 4, 5, 8, 10, 1], [1, 1, 1, 1, 2, 1, 2, 1, 1], [5, 1, 1, 1, 2, 1, 1, 1, 1], [10, 3, 5, 1, 10, 5, 3, 10, 2], [10, 8, 7, 4, 3, 10, 7, 9, 1], [5, 10, 8, 10, 8, 10, 3, 6, 3], [1, 1, 1, 1, 2, 1, 3, 1, 1], [3, 1, 1, 1, 2, 1, 1, 1, 1], [8, 10, 4, 4, 8, 10, 8, 2, 1], [2, 1, 1, 1, 2, 1, 1, 1, 1], [4, 1, 1, 1, 2, 1, 3, 1, 1], [4, 1, 1, 1, 2, 1, 1, 1, 1], [1, 1, 1, 1, 1, 1, 3, 1, 1], [3, 2, 2, 2, 2, 1, 3, 2, 1], [10, 10, 8, 10, 6, 5, 10, 3, 1], [3, 1, 1, 1, 2, 1, 1, 1, 1], [8, 3, 4, 9, 3, 10, 3, 3, 1], [7, 9, 4, 10, 10, 3, 5, 3, 3], [7, 1, 2, 3, 2, 1, 2, 1, 1], [5, 1, 1, 1, 2, 1, 3, 1, 1], [5, 1, 1, 1, 2, 1, 3, 1, 1], [5, 1, 1, 1, 2, 1, 1, 1, 1], [5, 8, 7, 7, 10, 10, 5, 7, 1], [4, 1, 2, 1, 2, 1, 2, 1, 1], [1, 1, 1, 1, 5, 1, 3, 1, 1], [5, 1, 2, 1, 2, 1, 3, 1, 1], [5, 3, 3, 2, 3, 1, 3, 1, 1], [10, 4, 5, 5, 5, 10, 4, 1, 1], [1, 1, 1, 1, 3, 2, 2, 1, 1], [1, 5, 8, 6, 5, 8, 7, 10, 1], [3, 1, 1, 1, 2, 1, 2, 2, 1], [7, 8, 7, 2, 4, 8, 3, 8, 2], [1, 1, 1, 1, 1, 1, 2, 1, 1], [5, 1, 1, 1, 1, 1, 3, 1, 1], [3, 3, 6, 4, 5, 8, 4, 4, 1], [5, 1, 2, 1, 2, 1, 2, 1, 1], [2, 1, 1, 1, 2, 5, 1, 1, 1], [5, 1, 2, 1, 2, 1, 3, 1, 1], [1, 1, 1, 1, 2, 1, 2, 1, 1], [3, 5, 7, 8, 8, 9, 7, 10, 7], [2, 1, 1, 1, 2, 1, 1, 1, 1], [4, 2, 1, 1, 2, 2, 3, 1, 1], [1, 1, 1, 1, 2, 1, 3, 1, 1], [9, 5, 5, 2, 2, 2, 5, 1, 1], [1, 1, 1, 1, 2, 1, 1, 1, 1], [5, 1, 4, 1, 2, 1, 3, 2, 1], [5, 7, 4, 1, 6, 1, 7, 10, 3], [8, 8, 9, 6, 6, 3, 10, 10, 1], [10, 10, 10, 4, 8, 1, 8, 10, 1], [8, 7, 4, 4, 5, 3, 5, 10, 1], [5, 3, 3, 1, 2, 1, 2, 1, 1], [4, 2, 3, 5, 3, 8, 7, 6, 1], [1, 4, 3, 10, 4, 10, 5, 6, 1], [4, 1, 1, 1, 2, 1, 3, 1, 1], [5, 7, 10, 10, 5, 10, 10, 10, 1], [10, 6, 6, 3, 4, 5, 3, 6, 1], [4, 1, 1, 1, 2, 1, 3, 1, 1], [5, 1, 3, 1, 2, 1, 2, 1, 1], [1, 1, 3, 1, 2, 1, 1, 1, 1], [1, 2, 3, 1, 2, 1, 3, 1, 1], [5, 1, 1, 1, 2, 1, 1, 1, 1], [3, 1, 1, 2, 2, 1, 1, 1, 1], [5, 5, 5, 6, 3, 10, 3, 1, 1], [6, 3, 3, 5, 3, 10, 3, 5, 3], [2, 1, 1, 1, 2, 1, 2, 1, 1], [5, 1, 1, 1, 2, 1, 2, 1, 1], [3, 1, 1, 1, 2, 1, 2, 1, 1], [7, 2, 4, 1, 3, 4, 3, 3, 1], [8, 4, 10, 5, 4, 4, 7, 10, 1], [4, 1, 3, 3, 2, 1, 1, 1, 1], [7, 2, 4, 1, 6, 10, 5, 4, 3], [5, 8, 8, 8, 5, 10, 7, 8, 1], [2, 3, 4, 4, 2, 5, 2, 5, 1], [4, 1, 1, 1, 2, 1, 1, 2, 1], [1, 1, 1, 1, 2, 1, 1, 1, 1], [4, 5, 5, 10, 4, 10, 7, 5, 8], [10, 6, 4, 1, 3, 4, 3, 2, 3], [3, 2, 2, 2, 2, 1, 4, 2, 1], [1, 1, 1, 1, 2, 1, 1, 1, 8], [1, 1, 1, 1, 1, 1, 3, 1, 1], [2, 1, 1, 2, 3, 1, 2, 1, 1], [3, 1, 2, 1, 2, 1, 3, 1, 1], [4, 1, 1, 1, 2, 1, 2, 1, 1], [8, 10, 3, 2, 6, 4, 3, 10, 1], [1, 1, 1, 1, 2, 1, 2, 1, 1], [4, 1, 3, 1, 2, 1, 2, 1, 1], [10, 10, 10, 10, 6, 10, 8, 1, 5], [1, 3, 3, 2, 2, 1, 7, 2, 1], [2, 1, 1, 1, 1, 1, 2, 1, 1], [10, 10, 10, 10, 10, 10, 4, 10, 10], [4, 2, 1, 1, 2, 1, 2, 1, 1], [5, 1, 1, 2, 2, 1, 2, 1, 1], [1, 1, 1, 1, 2, 1, 1, 1, 1], [3, 1, 1, 1, 2, 1, 2, 1, 1], [10, 9, 7, 3, 4, 2, 7, 7, 1]])
assert_equal(test_X.tolist(), [[1, 1, 1, 1, 2, 5, 1, 1, 1], [3, 1, 1, 1, 2, 1, 2, 1, 1], [5, 5, 5, 2, 5, 10, 4, 3, 1], [4, 7, 8, 3, 4, 10, 9, 1, 1], [1, 1, 1, 1, 2, 1, 1, 1, 1], [3, 3, 1, 1, 2, 1, 1, 1, 1], [5, 3, 1, 2, 2, 1, 2, 1, 1], [8, 10, 10, 10, 7, 5, 4, 8, 7], [1, 1, 1, 3, 1, 3, 1, 1, 1], [1, 1, 1, 1, 2, 1, 2, 1, 1], [8, 7, 5, 10, 7, 9, 5, 5, 4], [4, 1, 1, 1, 2, 1, 3, 2, 1], [10, 5, 6, 10, 6, 10, 7, 7, 10], [3, 1, 1, 1, 2, 1, 2, 1, 1], [5, 1, 1, 6, 3, 1, 1, 1, 1], [3, 1, 1, 3, 8, 1, 5, 8, 1], [9, 8, 8, 5, 6, 2, 4, 10, 4], [6, 10, 10, 2, 8, 10, 7, 3, 3], [7, 6, 4, 8, 10, 10, 9, 5, 3], [5, 1, 3, 1, 2, 1, 2, 1, 1], [3, 1, 1, 1, 2, 1, 2, 1, 1], [4, 1, 1, 2, 2, 1, 2, 1, 1], [3, 4, 5, 2, 6, 8, 4, 1, 1], [3, 1, 1, 1, 2, 1, 1, 1, 1], [10, 4, 6, 4, 5, 10, 7, 1, 1], [6, 6, 6, 5, 4, 10, 7, 6, 2], [5, 3, 1, 1, 2, 1, 1, 1, 1], [2, 1, 1, 1, 2, 1, 1, 1, 5], [5, 1, 1, 1, 2, 1, 1, 1, 1], [10, 7, 7, 3, 8, 5, 7, 4, 3], [1, 1, 1, 1, 2, 1, 1, 1, 1], [5, 10, 10, 10, 10, 10, 10, 1, 1], [8, 10, 10, 10, 5, 10, 8, 10, 6], [3, 1, 1, 1, 2, 1, 3, 1, 1], [4, 1, 1, 1, 2, 1, 1, 1, 1], [3, 1, 1, 1, 2, 1, 3, 1, 1], [3, 6, 6, 6, 5, 10, 6, 8, 3], [8, 5, 5, 5, 2, 10, 4, 3, 1], [4, 1, 1, 1, 2, 1, 1, 1, 1], [5, 4, 6, 7, 9, 7, 8, 10, 1], [3, 1, 1, 1, 2, 3, 3, 1, 1], [1, 1, 1, 1, 2, 5, 5, 1, 1], [1, 2, 3, 1, 2, 1, 2, 1, 1], [3, 1, 1, 1, 1, 1, 1, 1, 1], [1, 1, 1, 1, 2, 1, 1, 1, 1], [3, 1, 1, 1, 3, 2, 1, 1, 1], [1, 2, 2, 1, 2, 1, 1, 1, 1], [10, 6, 5, 8, 5, 10, 8, 6, 1], [5, 1, 1, 3, 2, 1, 1, 1, 1], [5, 1, 1, 1, 2, 1, 1, 1, 1], [10, 10, 9, 3, 7, 5, 3, 5, 1], [4, 2, 2, 1, 2, 1, 2, 1, 1], [10, 3, 6, 2, 3, 5, 4, 10, 2], [5, 1, 1, 2, 1, 1, 2, 1, 1], [5, 1, 3, 1, 2, 1, 2, 1, 1], [3, 1, 1, 2, 3, 4, 1, 1, 1], [10, 4, 4, 10, 2, 10, 5, 3, 3], [5, 4, 4, 5, 7, 10, 3, 2, 1], [5, 2, 2, 2, 2, 2, 3, 2, 2], [10, 8, 10, 1, 3, 10, 5, 1, 1], [1, 1, 1, 1, 2, 1, 3, 1, 1], [1, 1, 1, 1, 2, 2, 1, 1, 1], [2, 1, 1, 1, 2, 1, 3, 1, 1], [3, 1, 1, 1, 2, 1, 1, 1, 1], [3, 1, 1, 1, 2, 1, 3, 1, 1], [1, 1, 1, 1, 2, 1, 3, 1, 1], [1, 1, 1, 1, 2, 1, 3, 1, 1], [3, 1, 1, 1, 2, 1, 3, 1, 1], [5, 4, 4, 9, 2, 10, 5, 6, 1], [9, 5, 8, 1, 2, 3, 2, 1, 5], [1, 1, 2, 1, 2, 1, 2, 1, 1], [3, 1, 1, 1, 2, 1, 2, 1, 1], [1, 1, 1, 1, 2, 1, 2, 1, 1], [5, 1, 1, 1, 1, 1, 1, 1, 1], [5, 2, 2, 2, 3, 1, 1, 3, 1], [4, 1, 1, 1, 2, 1, 3, 1, 1], [10, 5, 5, 6, 8, 8, 7, 1, 1], [4, 1, 1, 1, 2, 1, 1, 1, 1], [3, 1, 1, 1, 2, 1, 2, 1, 1], [1, 1, 1, 1, 2, 1, 2, 1, 1], [5, 10, 10, 10, 6, 10, 6, 5, 2], [4, 1, 1, 3, 2, 1, 1, 1, 1], [10, 4, 3, 10, 3, 10, 7, 1, 2], [1, 1, 1, 1, 2, 1, 3, 1, 1], [1, 1, 1, 1, 2, 1, 1, 1, 1], [6, 10, 10, 10, 10, 10, 8, 10, 10], [6, 1, 3, 2, 2, 1, 1, 1, 1], [5, 3, 4, 3, 4, 5, 4, 7, 1], [10, 5, 8, 10, 3, 10, 5, 1, 3], [5, 3, 4, 1, 4, 1, 3, 1, 1], [10, 7, 7, 6, 4, 10, 4, 1, 2], [3, 1, 1, 1, 1, 1, 2, 1, 1], [10, 4, 7, 2, 2, 8, 6, 1, 1], [8, 2, 4, 1, 5, 1, 5, 4, 4], [5, 3, 3, 4, 2, 4, 3, 4, 1], [3, 4, 5, 3, 7, 3, 4, 6, 1], [5, 8, 9, 4, 3, 10, 7, 1, 1], [6, 10, 10, 2, 8, 10, 7, 3, 3], [7, 3, 4, 4, 3, 3, 3, 2, 7], [5, 1, 1, 1, 2, 1, 3, 1, 1], [6, 1, 1, 1, 2, 1, 2, 1, 1], [1, 1, 1, 2, 2, 1, 3, 1, 1], [10, 6, 6, 2, 4, 10, 9, 7, 1], [10, 10, 6, 3, 3, 10, 4, 3, 2], [4, 1, 1, 1, 2, 1, 3, 1, 1], [1, 1, 1, 3, 2, 3, 1, 1, 1], [2, 7, 10, 10, 7, 10, 4, 9, 4], [10, 2, 2, 1, 2, 6, 1, 1, 2], [3, 1, 1, 1, 1, 1, 2, 1, 1], [1, 1, 1, 1, 2, 1, 1, 1, 1], [7, 8, 3, 7, 4, 5, 7, 8, 2], [6, 1, 1, 3, 2, 1, 1, 1, 1], [3, 1, 1, 1, 2, 1, 2, 1, 1], [8, 10, 10, 8, 5, 10, 7, 8, 1], [1, 1, 1, 1, 1, 1, 3, 1, 1], [4, 3, 1, 1, 2, 1, 4, 8, 1], [5, 1, 1, 1, 2, 1, 2, 2, 1], [3, 4, 4, 10, 5, 1, 3, 3, 1], [1, 1, 1, 1, 2, 1, 2, 1, 1], [2, 1, 1, 1, 2, 1, 1, 1, 1], [1, 1, 1, 1, 2, 1, 2, 1, 1], [5, 4, 6, 6, 4, 10, 4, 3, 1], [3, 1, 1, 1, 2, 1, 2, 1, 1], [1, 1, 1, 1, 2, 1, 2, 1, 1], [5, 10, 10, 10, 4, 10, 5, 6, 3], [6, 10, 10, 10, 8, 10, 7, 10, 7], [1, 1, 1, 1, 2, 5, 1, 1, 1], [8, 4, 4, 1, 2, 9, 3, 3, 1], [1, 1, 1, 1, 2, 1, 3, 1, 1], [5, 10, 10, 8, 5, 5, 7, 10, 1], [1, 1, 1, 1, 10, 1, 1, 1, 1], [1, 1, 1, 1, 1, 1, 3, 1, 1], [10, 5, 5, 3, 6, 7, 7, 10, 1], [1, 1, 2, 1, 2, 1, 2, 1, 1], [4, 1, 1, 1, 3, 1, 1, 1, 1], [4, 10, 4, 7, 3, 10, 9, 10, 1], [2, 1, 1, 1, 3, 1, 2, 1, 1]])
assert_equal(test_y.tolist(), [0, 0, 1, 1, 0, 0, 0, 1, 0, 0, 1, 0, 1, 0, 0, 0, 1, 1, 1, 0, 0, 0, 1, 0, 1, 1, 0, 0, 0, 1, 0, 1, 1, 0, 0, 0, 1, 1, 0, 1, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 1, 0, 1, 0, 0, 0, 1, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 1, 1, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 1, 0, 1, 0, 0, 1, 0, 0, 1, 0, 1, 0, 1, 1, 1, 0, 1, 1, 1, 0, 0, 0, 1, 1, 0, 0, 1, 1, 0, 0, 1, 0, 0, 1, 0, 0, 0, 1, 0, 0, 0, 1, 0, 0, 1, 1, 0, 1, 0, 1, 0, 0, 1, 0, 0, 1, 0])
assert_equal(train_y.tolist(), [1, 0, 0, 0, 1, 0, 0, 0, 0, 1, 0, 1, 0, 1, 1, 0, 0, 0, 0, 1, 1, 0, 1, 1, 1, 1, 1, 1, 0, 0, 0, 0, 0, 0, 1, 1, 1, 0, 0, 0, 1, 0, 0, 1, 1, 0, 1, 0, 0, 1, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 1, 1, 0, 1, 0, 0, 0, 1, 0, 0, 0, 0, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 0, 1, 0, 1, 1, 0, 0, 0, 1, 0, 1, 0, 1, 0, 0, 0, 0, 0, 1, 0, 0, 1, 0, 0, 1, 0, 0, 0, 0, 0, 1, 0, 0, 1, 0, 1, 0, 0, 1, 1, 1, 0, 0, 0, 0, 0, 0, 1, 1, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 1, 1, 0, 1, 0, 0, 0, 1, 0, 0, 0, 1, 1, 0, 1, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 1, 1, 1, 1, 0, 1, 0, 1, 0, 1, 1, 1, 0, 0, 1, 0, 0, 0, 0, 1, 1, 0, 0, 0, 1, 0, 0, 1, 0, 0, 0, 0, 1, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 1, 0, 0, 0, 0, 0, 0, 1, 1, 1, 1, 0, 0, 1, 0, 1, 0, 1, 0, 0, 0, 0, 1, 0, 1, 0, 0, 0, 0, 1, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 1, 0, 1, 0, 0, 0, 1, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 1, 0, 0, 1, 0, 0, 0, 0, 1, 1, 0, 0, 0, 0, 1, 0, 0, 1, 0, 0, 0, 0, 1, 1, 0, 1, 0, 1, 0, 0, 0, 1, 1, 1, 0, 0, 0, 0, 0, 0, 0, 0, 1, 1, 0, 0, 0, 0, 0, 0, 0, 1, 1, 0, 0, 0, 0, 1, 1, 1, 0, 1, 0, 1, 0, 0, 0, 0, 0, 1, 0, 1, 1, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 1, 0, 0, 1, 1, 1, 0, 1, 1, 1, 0, 0, 0, 1, 0, 1, 0, 0, 1, 0, 1, 1, 1, 0, 0, 0, 1, 0, 1, 1, 1, 0, 0, 0, 1, 0, 1, 0, 0, 0, 0, 1, 1, 0, 0, 0, 1, 1, 0, 0, 1, 0, 0, 0, 1, 1, 0, 0, 0, 1, 0, 0, 0, 1, 0, 1, 0, 0, 1, 0, 0, 0, 0, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 1, 0, 0, 1, 0, 0, 0, 0, 0, 1, 0, 1, 1, 0, 0, 0, 0, 1, 0, 0, 0, 0, 1, 0, 1, 0, 1, 0, 0, 1, 0, 0, 0, 0, 1, 0, 0, 0, 1, 0, 0, 1, 1, 1, 1, 0, 1, 1, 0, 1, 1, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 1, 1, 0, 1, 1, 1, 0, 0, 1, 1, 0, 0, 0, 0, 0, 0, 1, 0, 0, 1, 0, 0, 1, 0, 0, 0, 0, 1])
```

# Problem 4: Training a Multi-Layer Perceptron classifier
For this problem create a Multi-Layer Perceptron classifier using sci-kit learns implementation Set random_state to be 100. Fit this on the training set. Lastly, name this model *modelMLPc*

<!-- grade_id: cell-ca83da335cd76ada, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-0f1c27e05d57087c, points: 1, grade: true, solution: false -->
```python
assert_almost_equal(.9708, modelMLPc.score(test_X, test_y), places=2)
```

# Problem 5: Using Deep Learning for classification

For this problem you will use keras to perform deep learning:
- First  convert the training and testing labels using the to_categorical method. 
- Next Create Sequential model and save it in a variable called $\texttt{model}$
- Add a densely connected neural net layer with 2 units and the input shape should be the number of columns in the training set
- Add a softmax activation function 
- Now configure the model for training by using the compile function from the Sequential model object. Set the loss to be *categorical_crossentropy*, the optimizier should be RMSprop and set the metrics to parameter so that the accuracy can be returned.

get the shape of both dimensions for the training features 

<!-- grade_id: cell-68c07c40d9715ffd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

```python
# Fit model
batch_size = 25
num_epoch = 20
val_spl = 0.2

# Train model, use last part of training data to validate
history = model.fit(train_X, y_trainc,
                    batch_size=batch_size, epochs=num_epoch,
                    verbose=1, validation_split = val_spl)

score = model.evaluate(test_X, y_testc, verbose=0)
```

```python
print("Testing Accuracy:", round(score[1], 3))
```

<!-- grade_id: cell-f287354529cf639c, points: 1, grade: true, solution: false -->
```python
assert_equal(model.get_config()[1]['config']['activation'], 'softmax')

sol = {'activation': 'linear',
 'activity_regularizer': None,
 'batch_input_shape': (None, 9),
 'bias_constraint': None,
 'bias_initializer': {'class_name': 'Zeros', 'config': {}},
 'bias_regularizer': None,
 'dtype': 'float32',
 'kernel_constraint': None,
 'kernel_initializer': {'class_name': 'VarianceScaling',
  'config': {'distribution': 'uniform',
   'mode': 'fan_avg',
   'scale': 1.0,
   'seed': None}},
 'kernel_regularizer': None,
 'name': 'dense_1',
 'trainable': True,
 'units': 2,
 'use_bias': True}

vals = []
for ans_key in model.get_config()[0]['config']:
    if ans_key == 'name': pass 
    else: assert_equal(model.get_config()[0]['config'][ans_key], sol[ans_key])
```

# Problem 6: Create a more complicated model

As mentioned in the notes, the power of deep learning lies in being able to build complex models that capture details in the data. For this problem, we will modify the previous deep learning model by adding in more layers and changing the activiation function. Read the instructions carefully.

- First  convert the training and testing labels using the to_categorical method. 
- Next Create Sequential model, and save it in a variable called $\texttt{model}$
- Add a densely connected neural net layer with 20 units and the input shape should be the number of columns in the training set
- Add a softmax activation function 
- Add a densely connected neural net layer with 2 units and the input shape should be the number of columns in the training set
- Add a tanh activation function 
- Add a densely connected neural net layer with 2 units and the input shape should be the number of columns in the training set
- Add a softmax activation function 
- Now configure the model for training by using the compile function from the Sequential model object. Set the loss to be *categorical_crossentropy*, the optimizier should be RMSprop and set the metrics to parameter so that the accuracy can be returned.

```python
train_dim_one = train_X.shape[0]
train_dim_two = train_X.shape[1]
test_dim_one = test_X.shape[0]

random.seed(0)
```

<!-- grade_id: cell-1b400ee6d1249b00, points: 0, grade: false, solution: true -->
```python
random.seed(0)
# YOUR CODE HERE
```

```python
random.seed(0)

# Fit model
batch_size = 25
num_epoch = 20
val_spl = 0.2

# Train model, use last part of training data to validate
history = model.fit(train_X, y_trainc,
                    batch_size=batch_size, epochs=num_epoch,
                    verbose=1, validation_split = val_spl)

score = model.evaluate(test_X, y_testc, verbose=0)
```

Now let's see how accurate the model is

```python
print("Testing Accuracy:", round(score[1], 3))
```

<!-- grade_id: cell-db3bb6bed2946d73, points: 1, grade: true, solution: false -->
```python
assert_equal(model.get_config()[3]['config']['activation'], 'tanh')

sol = {'activation': 'linear',
 'activity_regularizer': None,
 'batch_input_shape': (None, 9),
 'bias_constraint': None,
 'bias_initializer': {'class_name': 'Zeros', 'config': {}},
 'bias_regularizer': None,
 'dtype': 'float32',
 'kernel_constraint': None,
 'kernel_initializer': {'class_name': 'VarianceScaling',
  'config': {'distribution': 'uniform',
   'mode': 'fan_avg',
   'scale': 1.0,
   'seed': None}},
 'kernel_regularizer': None,
 'name': 'dense_1',
 'trainable': True,
 'units': 20,
 'use_bias': True}

vals = []
for ans_key in model.get_config()[0]['config']:
    if ans_key == 'name': pass 
    else: assert_equal(model.get_config()[0]['config'][ans_key], sol[ans_key])
```

```python

```