# Module 1 Assignment

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

```python
from nose.tools import assert_equal, assert_almost_equal, assert_is_instance
import numpy as np
from pprint import pprint
from collections import Counter
import csv, json, helper
import xml.etree.ElementTree as ET
```

Load in the cars datafile

```python
data_file = "./imports-85.data"
```

# Problem 1: Write a function that reads in a text file and returns the column names.

Write a function called "get_names" which takes in a comma separate file as an input and outputs a list containing the data in the file, and also returns the header of the file, containing all the column names.

<!-- grade_id: cell-f7463603e5fea09c, points: 0, grade: false, solution: true -->
```python
def get_names(file):
    
    '''
    Inputs
    --------
    file: the file to be read
    
    Returns
    -------
    car_data: a list contaning all the data from the file
    header: the header of the datafile
 
    '''
    
    # YOUR CODE HERE
        
    return car_data, header
    
```

<!-- grade_id: cell-3ef73391403401dc, points: 1, grade: true, solution: false -->
```python
data, col_names = get_names(data_file)

true_names = ['symboling', 'normalized-losses', 'make',
              'fuel-type', 'aspiration', 'num-of-doors', 'body-style', 'drive-wheels', "'engine-location'",
              'wheel-base', 'length', 'width', 'height', 'curb-weight', 'engine-type', 'num-of-cylinders',
              'engine-size', 'fuel-system', 'bore', 'stroke', "compression-ratio'", 'horsepower', 'peak-rpm',
              'city-mpg', 'highway', 'price']


assert_equal(true_names,col_names)
assert_is_instance(data,list)
assert_equal(data[1][0],'3')
```

# Problem 2: Reading in JSON
For this problem write a function called read_json, that accepts 1 parameter. This parameter will contain the file path to a json dataset about projects that the world bank has funded. Your task is to:
- Read the first line in the dataset
- Use python's [json](https://docs.python.org/3/library/json.html) library to decode the json.
- Lastly return the decoded json (this should have the type of a dictionary if done correctly)

<!-- grade_id: cell-13521609cdadce80, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-2cc5041c33116c85, points: 1, grade: true, solution: false -->
```python
_file = './world_bank.json'
data = read_json(_file)
assert_equal(type(helper.json_l3), type(data), msg="You did not decode the data correctly")
assert_equal(helper.json_l3, data, "You data does not match the solutions")
```

```python
world_bank_names = []
```

# Problem 3: Grabbing project names from the world bank.
For this problem grab the names of each project funded by the world bank and append it to a list called world_bank_names. Remember that the data is stored in a variable called data.

<!-- grade_id: cell-17fa2dd7ddfeca0a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-afb5e85f6e7f0adc, points: 1, grade: true, solution: false -->
```python
assert_equal(world_bank_names, ['World Development Indicators',
 'Africa Development Indicators',
 'Global Economic Monitor',
 'Research Datasets and Analytical Tools',
 'Millennium Development Goals',
 'GovData360',
 'Data on Statistical Capacity',
 'Doing Business',
 'Education Statistics',
 'Enterprise Surveys'])
```

# Problem 4: Working with XML
For this problem write a function called book_xml that accepts 1 parameter. This parameter will contain the file path to a xml dataset about books. (I suggest you take a look at books.xml *be sure not to modify this file or you may fail the assertion test*). Your task is to:
- Read in the xml file using python's [ElementTree](https://docs.python.org/3/library/xml.etree.elementtree.html)s XML documentation.
- Grab the root node of the tree.
- Next iterate through the children of the root node and append the **title** of each book to a list. Return this list.

<!-- grade_id: cell-f401375361368483, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-e2d767c2d829b65e, points: 1, grade: true, solution: false -->
```python
book_titles = book_xml('books.xml')
assert_equal(book_titles, ["XML Developer's Guide", 'Midnight Rain', 'Maeve Ascendant', "Oberon's Legacy", 'The Sundered Grail', 'Lover Birds', 'Splish Splash', 'Creepy Crawlies', 'Paradox Lost', 'Microsoft .NET: The Programming Bible', 'MSXML3: A Comprehensive Guide', 'Visual Studio 7: A Comprehensive Guide'])
```

# Problem 5: Finding the most common author.
For this problem write a function called most_common_author that accepts 1 parameter. This parameter will contain the file path to a xml dataset about books. (I suggest you take a look at books.xml *be sure not to modify this file or you may fail the assertion test*). Your task is to:
- Read in the xml file using python's [ElementTree](https://docs.python.org/3/library/xml.etree.elementtree.html)s XML documentation.
- Grab the root node of the tree.
- Next iterate through the children of the root node and append the **author** of each book to a list.
- Now iterate through the list of the authors and return the most common author (there are a lot of ways that you can do this. I recommend taking a look at the [Counter documentation](https://docs.python.org/3/library/collections.html#counter-objects) in the collections library's doc). Be sure that this generalizes to other xml datasets that have the same structure.

<!-- grade_id: cell-ca7648ef58ace751, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-1d1f2b95e5d9f27d, points: 1, grade: true, solution: false -->
```python
com_auth_1 = most_common_author('books.xml')
assert_equal(com_auth_1, 'Corets, Eva')
com_auth_2 = most_common_author('books1.xml')
assert_equal(com_auth_2, 'Galos, Mike')

```

```python

```